Differential increase of hippocampal subfield volume after socio-affective mental training relates to reductions in diurnal cortisol

Author response:

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer #1 (Public Review):

Valk and Engert et al. examined the potential relations between three different mental training modules, hippocampal structure and functional connectivity, and cortisol levels (stress) over a 9-month period. They found that among the three types of mental training: Presence (attention and introspective awareness), Affect (socio-emotional - compassion and prosocial motivation), and Perspective (socio-cognitive - metacognition and perspective taking) modules; Affect training most robustly related to changes in hippocampal structure and function - specifically, CA1-3 subfields of the hippocampus. Moreover, change in intrinsic functional connectivity related to changes in diurnal cortisol release and long-term cortisol exposure. These changes are proposed to result from a combination of factors, which is supported by multivariate analyses showing changes across subfields and training content relate to cortisol changes.

The authors demonstrate that mindfulness training programs are a potential avenue for stress interventions that impact hippocampal structure and cortisol, providing a promising approach to improve health. The data contribute to the literature on plasticity of hippocampal subfields during adulthood, the impact of mental training interventions on the brain, and the link between CA1-3 and both short- and long-term stress changes.

The authors thoughtfully approached the study of hippocampal subfields, utilizing a method designed for T1w images that outperformed Freesurfer 5.3 and that produced comparable results to an earlier version of ASHS. The authors note the limitations of their approaches and provide detailed information on the data used and analyses conducted. The results provide a strong basis from which future studies can expand using computational approaches or more fine-grained investigations of the impact of mindfulness training on cortisol levels and the hippocampus.

We thank the Reviewer for the positive re-evaluation and summary of our findings and work. We made additional change as suggested and hope this clarified any open points.

I have a few additional suggestions. Clarifying the language around the multivariate results and the impact across subfields and training modules would be helpful.

We are happy to provide further clarifications with respect to the multivariate results and the impact of training on subfields.

The multivariate analyses served as a final step to explore any potential connections between training modules and hippocampal subfields, beyond just the link between CA1-3 and the Affect Module. These additional analyses were suggested by the Reviewers, and we, as authors, agreed that taking a broader view of how different parts of the hippocampus interact with overall changes can provide valuable insights into the relationship between mental training, cortisol fluctuations, and changes in CA1-3 subfields.

We employed a multivariate partial least squares method, which aims to identify the directions in the predictor space that account for the most variance in changes observed, by creating latent variables. Initially, we investigated whether there was a general connection between CA1-3 subfields and cortisol changes, regardless of which training module produced these effects. Our findings confirmed a consistent relationship across all three training modules, indicating a strong association between cortisol changes, particularly markers such as AUC and slope change, and alterations in CA1-3 structure and functional connectivity. We explored a model incorporating changes across all hippocampal subfields and stress markers across different modules. In the right hemisphere, changes in the volume of the CA1-3 subfield were more strongly associated with stress markers, compared to other subfields. However, this association was less pronounced in the left hemisphere.

Our multivariate approach captured fluctuations across subfields and modules beyond group-level associations, leading to a more nuanced interpretation. While the univariate analysis of module-specific changes in volume and associations within the Affect Module may offer a straightforward interpretation, as they coincide with increases in CA1-3 volume, the multivariate analysis also accounts for individual-level changes not observed at the group level using a data-driven approach. Overall these findings are in line with the group-level observations, yet provide nuance on specificity.

We clarified these considerations further in the manuscript;

Abstract:

“Notably, using a multivariate approach, we found that other subfields that did not show group-level changes also contributed to changes in cortisol levels.”

Results:

“We employed a multivariate partial least squares method, which aims to identify the directions in the predictor space that account for the most variance in changes observed, by creating latent variables. Initially, we investigated whether there was a general connection between CA1-3 subfields and cortisol changes, regardless of which training module produced these effects.”

Discussion:

“Finally, through conducting multivariate analysis, we once more noticed associations between changes in CA1-3 volume and functional adaptability and alterations in stress levels, particularly prominent within the Affect Module. Integrating all subfields into a unified model highlighted a distinct significance of CA1-3, although for the left hemisphere, we observed a more diverse range of contributions across subfields. In summary, we establish a connection between a socio-emotional behavioral intervention, shifts in hippocampal subfield structure and function, and decreases in cortisol levels among healthy adults.

Although the univariate examination of changes specific to modules in volume and connections within the Affect Module presents how changes in cortisol align with group-level rises in CA1-3 volume, the multivariate analysis extended this observation through considering individual-level alterations not discernible at the group level through a data-driven method. These results generally corresponded with observations at the group level but offer additional insights into specificity, and hint at system-level alterations.”

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: The hippocampus is a central modulator of the HPA-axis, impacting the regulation of stress on brain structure, function, and behavior. The current study assessed whether three different types of 3-months mental Training Modules geared towards nurturing a) attention-based mindfulness, b) socio-affective or c) socio-cognitive skills may impact hippocampal organization by reducing stress. We evaluated mental training-induced changes in hippocampal subfield volume and intrinsic functional connectivity, by combining longitudinal structural and resting-state fMRI connectivity analysis in 332 healthy adults. We related these changes to changes in diurnal and chronic cortisol levels. We observed increases in bilateral cornu ammonis volume (CA1-3) following the 3-months compassion-based module targeting socio-affective skills (Affect module), as compared to socio-cognitive skills (Perspective module) or a waitlist cohort with no training intervention. Structural changes were paralleled by relative increases in functional connectivity of CA1-3 when fostering socio-affective as compared to socio-cognitive skills. Furthermore, training-induced changes in CA1-3 structure and function consistently correlated with reductions in cortisol output. Notably, using a multivariate approach, we found that other subfields that did not show group-level changes also contributed to changes in cortisol levels, suggesting that circuit-level changes within the hippocampal formation are linked to diurnal stress markers. Overall, we provide a link between a socio-emotional behavioural intervention, changes in hippocampal subfield structure and function, and reductions in cortisol in healthy adults.

Revised version: The hippocampus is a central modulator of the HPA-axis, impacting the regulation of stress on brain structure, function, and behavior. The current study assessed whether three different types of 3-months mental Training Modules geared towards nurturing a) attention-based mindfulness, b) socio-affective or c) socio-cognitive skills may impact hippocampal organization by reducing stress. We evaluated mental training-induced changes in hippocampal subfield volume and intrinsic functional connectivity, by combining longitudinal structural and resting-state fMRI connectivity analysis in 332 healthy adults. We related these changes to changes in diurnal and chronic cortisol levels. We observed increases in bilateral cornu ammonis volume (CA1-3) following the 3-months compassion-based module targeting socio-affective skills (Affect module), as compared to socio-cognitive skills (Perspective module) or a waitlist cohort with no training intervention. Structural changes were paralleled by relative increases in functional connectivity of CA1-3 when fostering socio-affective as compared to socio-cognitive skills. Furthermore, training-induced changes in CA1-3 structure and function consistently correlated with reductions in cortisol output. Notably, using a multivariate approach, we found that other subfields that did not show group-level changes also contributed to changes in cortisol levels. Overall, we provide a link between a socio-emotional behavioural intervention, changes in hippocampal subfield structure and function, and reductions in cortisol in healthy adults.

Introduction

Stress-related disorders rank among the leading causes for disease burden world-wide (1), and the global stress load has increased even more dramatically in recent years (2, 3). It is therefore essential to find ways to efficiently prevent or reduce stress (4). In recent years, research has shown that contemplative mental training programs can be efficient in stress reduction ((5–7); for a meta-analysis see (8)), while simultaneously inducing brain plasticity (9–11). It is, however, still unclear which types of mental practices are most efficient in reducing stress and inducing stress-related brain plasticity. Furthermore, stress is a multi-layered construct (12), and most studies focused on stress-related self-reports and questionnaires (8). A less investigated marker in the stress reduction context through contemplative mental training is diurnal cortisol, from which summary indices such as the cortisol awakening response (CAR), the total diurnal output and the diurnal cortisol slope are frequently investigated (13). The steroid hormone and glucocorticoid cortisol is the end-product of the hypothalamic-pituitary-adrenal (HPA) axis and a key player in stress regulation (for reviews see e.g.,(14, 15)). Moreover, cortisol is considered an important mediator of the relation between chronic stress and stress-related disease (16, 17). Previous research suggests an association between hippocampal structural integrity and stress related cortisol activity (18, 19), although findings are inconclusive. To close these gaps, we here investigate the differential efficiency of three types of mental training (attention-based, socio-affective and socio-cognitive) on their ability to induce structural as well as functional plasticity of hippocampal subfields and reduce diurnal cortisol levels.

The hippocampus has a high glucocorticoid receptor density (20–23), making this region a target of investigations into stress-related brain changes. Having a three layered allocortex, the hippocampal formation consists of multiple subfields, or zones, starting at the subiculum (SUB) and moving inward to the hippocampus proper; the cornu ammonis (CA1-3), and dentate gyrus (CA4/DG)(24–27). These subfields have unique microstructure (24–28) and participate differently in the hippocampal circuitry (29), likely implicating different functional contributions (30–35). Indeed, intrinsic functional MRI analyses have shown that the hippocampal subfields show functional correlation with a broad range of cortical regions, part of visual, control, and default functional networks (28, 35–38). Hippocampal subfield volumes and associated intrinsic functional connectivity have been shown to be heritable (36, 39), indicating that individual variation in subfield structure and function is, in part, under genetic control. Other lines of research have reported hippocampal structure and function to be highly sensitive to contextual factors, such as stress (23). Mediated through its dense network of glucocorticoid receptors, the hippocampus transmits the negative feedback signals of a wide range of glucocorticoid levels on HPA axis activity (22). Through this inhibitory role on HPA axis dynamics, it is linked to emotional reactivity (40), stress sensitivity (19, 41–43), and causally involved in a variety of stress-related disorders (44).

Previous brain imaging research has examined the relationship between cortisol activity and hippocampal structure and function. Most of this research measured saliva cortisol levels to gauge the diurnal cortisol profile. Thus, a reduced cortisol awakening response, the response to the anticipated demands of the upcoming day (45), has been associated with smaller hippocampal volume in healthy individuals (46–48) and different psychiatric (49, 50) and metabolic (48, 51) conditions. In fact, the examination of patients with temporal lobe damage suggested that hippocampal integrity may be a necessary condition for the proper mounting of the CAR (52, 53). Next, to changes in hippocampal structure, alterations in hippocampal functional connectivity have been reported to be associated with changes in cortisol levels (54, 55). There is also contrary work showing associations between elevated awakening, evening, diurnal, or 24-hour cortisol levels in healthy elderly with age-related hippocampal atrophy (56–59) and, again, samples with psychiatric conditions (60, 61). While such inconsistencies in previous neuroimaging work may reflect the fact that different indices of diurnal cortisol tap different facets of HPA axis regulation, the samples studied have been diverse in terms of health status, small in size, and largely cross-sectional. In addition, associations between stress and hippocampal structure and function over time are incompletely understood. Thus, longitudinal studies, such as mental training studies aimed at stress reduction that repeatedly assess both brain and cortisol release, may help to better understand the dynamic relationships between stress, cortisol, and hippocampal structure and function.

In recent years, contemplative mental training interventions, such as the mindfulness-based stress reduction (MBSR) program (62) or compassion-focused therapy (63), have gained in popularity as potential therapeutic tools to improve mental and physical health (64) and reducing stress (8). These mental training interventions can have a positive impact on the practitioner’s stress sensitivity, making them a suitable model to investigate the interrelationship between training-related changes in hippocampal structure, function, and cortisol output. Next to reductions in reactive measures following acute psychosocial stress induction in the laboratory (7), reduced subjective-psychological stress load is the most widely reported outcome (for a review, see (8)). Evidence for lower diurnal cortisol output stems mainly from mindfulness-based interventions, notably MBSR, for which reductions in CAR and afternoon/evening cortisol levels have been reported in healthy and diseased individuals (65–67). Moreover, other work in the current sample has shown that hair cortisol and cortisone are reduced through mental practice (5). Hair cortisol measurements have been suggested to provide a window into long-term impact of cortisol exposure (68). These findings are contrasted by numerous null results (for meta-analyses see: (69, 70)), possibly due to modest samples sizes and mixed effects of different training contents on stress-related processes. Furthermore, 8-weeks mindfulness programs such as MBSR and others typically cultivate different types of mental practices, making it difficult to understand which type of mental practice is most efficient in reducing different types of outcomes, including various stress-markers (see also (5–7)).

The current study, therefore, investigated differential effects of distinct mental training practices onto the association between changes in hippocampal subfields and underlying stress-related diurnal cortisol profiles changes in the context of a large-scale 9-month mental training study, the ReSource Project (71). We explored impact of long-term exposure to stress onto hippocampal subfields as a function of mental training in a subset of individuals (5). Healthy participants attended three 3-months Training Modules termed Presence (cultivating attention and interoceptive awareness), Affect (cultivating compassion, prosocial motivation and dealing with difficult emotions) and Perspective (cultivating metacognition and perspective-taking on self and others) (Figure 1). Presence resembles typical mindfulness-based interventions, but excludes socio-emotional or socio-cognitive practices (62, 72). By contrast, Affect and Perspective target social skills through the training of either socio-emotional and motivational skills such as empathy, compassion and care (Affect) or socio-cognitive skills such as perspective taking on self and others (Perspective). In previous work, stemming from the same participant sample as examined here, we found a reduction in CAR specifically after the training of socio-affective capacities(6), and of acute stress reactivity after the training of socio-affective or so-ciocognitive capacities (7). Alternatively, but also in the current sample, different types of mental practices equally reduced hair cortisol levels, a marker of long-term stress (5). This suggests that the content of mental training has a specific effect on daily cortisol changes but not on long-term stress levels. Our group could also show differentiable training-related changes in cortical structure and intrinsic functional organization following the three ReSource project Training Modules, illustrating the existence of training-related structural plasticity of the social brain (9, 73). Domain-specific changes in hippocampal subfield structure and intrinsic functional connectivity, and how these relate to mental training specific changes in stress-related diurnal cortisol output, have not yet been studied. We, therefore, examined whether module-specific changes in diurnal cortisol levels may relate to specific structural and intrinsic functional changes in different hippocampal subfields and functional resting state data.

We evaluated the longitudinal relationship between hippocampal subfield volumetry, a quantitative index of hippocampal grey matter, and studied whether volumetric changes were paralleled by subfields’ resting-state functional connectivity in a large sample of healthy adults participating in the ReSource Project (71). This enabled us to evaluate training effects on hippocampal structure, function, and their associations with cortisol as a function of mental training targeting either attention-based mindfulness (Presence), socio-affective (Affect), or socio-cognitive (Perspective) skills. Hippocampal structure was quantified via a surface-based multi-template algorithm that has been shown to perform with excellent accuracy in healthy and diseased populations of a comparable age range as the currently evaluated cohort (74). Such a model is good to represent different subfields in vivo, which have a differentiable structure and function (75, 76), and thus may show differentiable changes as a function of mental training. We expect that assessment of hippocampal sub-regions may help to accurately map circuit plasticity as a result of potential stress reduction, and to observe that changes in hippocampal structure are paralleled by changes in functional connectivity of hippocampal subfield functional networks. To model the interplay between individual-level correspondence in hippocampal and stress markers, we assessed the association of changes in hippocampal structure and function with changes in several indices of diurnal cortisol release.

Results

We analyzed structural MRI, resting-state functional MRI, as well as cortisol-based stress markers from the large-scale ReSource Project (71). For details, see http://resource-project.org and the preregistered trial https://clinicaltrials.gov/ct2/show/NCT01833104.

In the Resource study, participants were randomly assigned to two training cohorts (TC1, N=80; TC2, N=81) and underwent a 9-months training consisting of three sequential Training Modules (Presence, Affect, and Perspective) with weekly group sessions and daily exercises, completed via cell-phone and internet platforms (Figure 1, Table 6-8, Materials and Methods and Supplementary Materials for more details). TC1 and TC2 started their training regimen with the Presence module, and then underwent the latter two modules in different orders (TC1: Affect-Perspective; TC2 Perspective-Affect) to serve as active control groups for each other (Figure 1C). Another active control group (TC3; N=81) completed three months of Affect training only. Additionally, a matched test-retest control cohort did not undergo any training (RCC, N=90). All participants were examined at the end of each 3-months module (T1, T2, T3) using 3T MRI, behavioral and peripheral-physiological measures that were identical to the baseline (T0) measures.

Change in bilateral CA1-3 volume following Affect mental training

The above design allowed us to examine whether the volume of hippocampal subfields shows increases or decreases following the distinct Training Modules. We tracked longitudinal changes in hippocampal subfield volumes using mixed-effects models (74). Excluding participants with missing or low quality structural and functional data, the sample included 86 individuals for Presence, 92 individuals for Affect, 83 individuals for Perspective, and 61 active controls (Affect) with hippocampal change scores. We included 164 change scores of retest controls over T1, T2, T3. To study whether there was any training module-specific change in hippocampal subfield volumes following mental training, we compared training effects between all three Training Modules (Presence, Affect, and Perspective). Main contrasts were: Presence vs Active control (between subjects) and Affect vs Perspective (within subjects). Supplementary comparisons were made vs retest controls and within training groups. We observed relative increases in right cornu ammonis 1-3 (CA1-3), but not in subiculum (SUB) or CA4 and dentate gyrus (CA4/DG) subfields, following Affect versus Perspective training (left: t=2.360, p=0.019, FDRq(q)>0.1, Cohens D=0.282; right: t=2.930, p=0.004, q=0.022, Cohens D =0.350), that could be attributed to subtle increases (p<0.05) in bilateral CA1-3 volume following Affect (left: t=2.495, p=0.013, q=0.08, M: 25.511, std: 130.470, CI [-1.509 52.531; right: t=2.374, p=0.018, q>0.1, M: 40.120, std: 181.300, CI [2.573 77.666]), and subtle decreases (p<.05) in right CA1-3 volume following Perspective (left: t=-1.143, p>0.1, q>0.1, M:-23.048, std: 137.810, CI [-53.139 7.043; right: t=-2.118, p=0.035, q>0.1, M:-39.602, std: 208.470, CI [-85.122 5.917]). We did not observe differences between Presence and the Active control cohort, Affect TC3. Overall, for all hippocampal subfields, findings associated with volume increases in CA1-3 following the Affect training were most consistent across timepoints and contrasts (Supplementary File 1a-g). Moreover, associations between CA1-3 and Affect, relative to Perspective, seemed to go largely above and beyond changes in the other subfields (left: t-value: 2.298, p=0.022, Q>0.1; right: t-value: 3.045, p=0.0025, Q=0.015, see further Supplementary File h). We observed no overall change in hippocampal subfield volume following mental training of nine months (Supplementary File 1i). Although stereotaxic normalization to MNI space would in theory account for global sex differences in intra-cranial volume (ICV), we still observed sex differences in various subfield volumes at baseline. Yet, accounting for ICV did not impact our main results, suggesting changes in CA1-3 following Affect were robust to sex differences in overall brain volume (Supplementary File 1j).

Training induced plasticity of hippocampal subfield volume.

A) Training modules; B) Training design; C) Subfield volumes in left and right hemispheres across individuals and timepoints; D) Scatterplot of subfield volumes as a function of timepoints and training cohorts.

Increased functional connectivity of CA1-3 following socio-affective versus socio-cognitive mental training

Subsequently, we studied whether volumetric change in hippocampal CA1-3 would show corresponding changes in intrinsic function following the Affect mental training. To probe the CA1-3 functional connectivity networks per subfield, we mapped the top 10% of normalized functional connections at baseline. Functional connectivity was strongest to medial prefrontal regions, precuneus extending to posterior cingulate, anterior temporal regions and angular gyrus (CA1-3: Figure 2; see Supplementary Materials for other subfields). Evaluating functional connectivity changes, we found that the right CA1-3 functional network showed differential changes when comparing Affect training to Perspective training (2.420, p=0.016, q=0.032, Cohens D =0.289), which could be attributed to subtle (p<.05) decreases in right CA1-3 mean FC following Perspective (t=-2.012, p=0.045, q>0.1, M:-0.024, std: 0.081, CI [-0.041 -0.006]), but not Affect training (t=1.691, p=0.092, q>0.1, M: 0.010, std: 0.098, CI [-0.01 0.031]); changes were not present when comparing Affect training versus retest control (Table 1 and Supplementary File 1k-q). Comparing Affect TC3 relative to Presence training, we did not observe changes (Table 1). No other subfield showed differential changes in main contrasts within its functional network.

Exploring whether particular regions within the CA1-3 network showed alterations in intrinsic functional connectivity when comparing Affect to Perspective, we investigated connectivity changes within regions of the subfields’ functional networks. Left CA1-3 connectivity showed decreases in connectivity to left posterior insula when comparing Affect to Perspective training (FDRq<0.05; t=-3.097, p=0.003, Cohens D=-0.370). On the other hand, we observed connectivity increases between right CA1-3 to right mPFC for the same contrast (FDRq<0.05; t=3.262, p=0.002, Cohens D =0.389). No other subfield’s functional connectivity showed alterations when comparing Affect to Perspective or Presence to Affect TC3. These analyses indicate an overlap between volumetric increases and functional alterations when comparing changes following socio-affective mental training in CA1-3. In particular, the moderately consistent CA1-3 volume increases following Affect training were complemented with differential functional connectivity alterations of this subfield when comparing Affect to Perspective training.

Training induced plasticity of CA1-3 functional connectivity.

A) upper: CA1-3 functional connectivity at baseline, top 10% of regions representing the CA1-3 functional network; lower: scatter plot visualizing change within the CA1-3 network across timepoints and groups; networks and scatters of SUB and CA4/DG are available in the supplements; B) Regional change within CA1-3 functional network Affect versus Perspective (FDRq<0.05); right: scatter plot visualizing mean change within the CA1-3, FDRq<0.05 regions across timepoints and groups.

Association between change in subfield volume, function, and stress markers

Last, we probed whether group-level changes in hippocampal subfield CA1-3 volume would correlate with individual-level changes in diurnal cortisol indices (Presence: n= 86; Affect: n=92; Perspective: n=81), given that the hippocampal formation is a nexus of the HPA-axis (19). We took a two-step approach. First, we studied univariate associations between cortisol and subfield change, particularly focusing on the Affect module and CA1-3 volume based on increases in CA1-3 volume identified in our group-level analysis. We observed that increases in bilateral CA1-3 following Affect showed a negative association with change in total diurnal cortisol output (operationalized as the area under the curve with respect to ground; AUCg) (left: t= -2.237, p=0.028, q=0.056; right: t=-2.283, p=0.025, q=0.05), indicating that with a reduction in stress-levels as measured by AUCg, there were increases in CA1-3 volume. Post-hoc analyses indicated no other subfield showed an association with AUCg, or with any of the other cortisol indices, below p<0.05 (Supplementary File 1r). Assessing the associations between cortisol indices and the right CA1-3 subfield functional networks in Affect (n=92), we could not observe individual level modulation of diurnal cortisol markers and group-level effects (right CA1-3 functional network change and cortisol markers or within the PI or mPFC ROI, Table 3 and Supplementary File 1s). Yet, we observed positive associations between mean functional network of left CA1-3 and diurnal slope (t=2.653, p=0.01, q=0.02) and AUCg (t=2.261, p=0.026, q=0.052), Supplementary File 1t. When assessing whether particular regions within the CA1-3 network showed alterations in intrinsic functional connectivity, we observed that AUCg modulated increases in connectivity between left CA1-3 and parietal occipital area (q<0.05). These analyses extend group-level observations regarding the relation between socio-affective mental training and CA1-3 structure to the individual-level. Again, we observed some consistency in structure and function in case of CA1-3. We did not observe alterations in CA1-3 volume in relation to change in cortisol markers in Presence or Perspective. Yet, for Presence we observed association between slope and LCA4/DG change (t=-2.89, p=0.005, q=0.03), (Supplementary Tables File 1uv). In case of intrinsic function, we also did not observe alterations in CA1-3 in relation to change in cortisol markers in Presence or Perspective, nor in other subfields (Supplementary File 1wx). When evaluating overall associations between diurnal cortisol change in CA1-3 volume in all modules combined, (Presence, Affect, and Perspective), we observed comparable patterns as for Affect only, further underscoring the association between cortisol markers and CA1-3 (Figure 3B; Supplementary File 1y and z). Last, we explored whether associations of subfield volume were found with levels of hair cortisol, a long-term marker of systemic cortisol exposure, in a sub-sample of N=44 participants repeatedly tested across modules (Presence, Affect, and Perspective), based on previous observations of domain-general effects of mental training on cortisol and cortisone (5). We identified consistent associations between increases in LCA1-3 volume and intrinsic function and hair cortisol decreases (volume: t=-2.574, p=0.011, q=0.022, function: t=-2.700, p=0.008, q=0.016). Exploring effects in subfields other than CA1-3 we revealed associations between right CA4/DG volume and cortisol (t=-3.138, p=0.002, q = 0.01) and left SUB function (t=-2.890, p=0.005, q = 0.03) (Figure 3B; Supplementary File 1za and zb).

Associations between changes in structure and function of hippocampal subfield volume and markers of stress change.

A). Upper left: Correlation between hippocampal subfield volume change in Affect and CAR, slope, and AUC markers of stress change, Upper right: Correlation between hippocampal subfield intrinsic functional change in Affect and CAR, slope, and AUC markers of stress change, middle: Scatter plots visualize the correlation between volume change and cortisol marker change (below p<0.05), bottom: region level change within left CA1-3, FDRq<0.05. CA1-3 is the focus of this analysis based on our group-level findings and highlighted with boxes in A; B). Upper: Overall impact of diurnal cortisol markers on hippocampal subfield volume and function over Presence, Affect and Perspective; Lower: Overall impact of hair cortisol markers on hippocampal subfield volume and function over Presence, Affect and Perspective

Earlier version: Second, to further explore the question whether there is concordant cortisol-related structural and functional plasticity in the hippocampus we adopted a multivariate partial least square approach, with 1000 permutations to account for stability (77, 78) and bootstrapping (100 times) with replacement. We ran two complementary models including all Training Modules whilst regressing out age, sex and random effects of subject. First, we focused on the bilateral CA1-3, as this region showed increases in volume following Affect training and differential change between Affect and Perspective training in our resting state analyses. In the second model included structural and functional data of all subfields. Both models included all stress markers. We found that both models could identify significant associations between cortisol stress markers and hippocampal plasticity (FDRq<0.05), and that in particular Affect showed strongest associations with the latent markers for CA1-3 (Table 5). Both analyses showed inverse effects of subfield structure and function in relation to stress markers and both slope and AUC changes showed strongest associations with the latent factor.

Revised version: We employed a multivariate partial least squares method, with 1000 permutations to account for stability (77, 78) and bootstrapping (100 times) with replacement, which aims to identify the directions in the predictor space that account for the most variance in changes observed, by creating latent variables. Initially, we investigated whether there was a general connection between CA1-3 subfields and cortisol changes, regardless of which Training Module produced these effects. This analysis was motivated by our observations that the bilateral CA1-3 showed increases in volume following Affect training and differential change between Affect and Perspective training in our resting state analyses. In a second model included structural and functional data of all subfields. Both models included all stress markers, and we regressed out age, sex and random effects of subject. We found that both models could identify significant associations between cortisol stress markers and hippocampal plasticity (FDRq<0.05), and that in particular Affect showed strongest associations with the latent markers for CA1-3 (Table 5). Both analyses showed inverse effects of subfield structure and function in relation to stress markers and both slope and AUC changes showed strongest associations with the latent factor.

Multivariate associations between changes in structure and function of hippocampal subfield volume and markers of stress change in Affect.

A). Multivariate associations between bilateral CA1-3 volume and intrinsic function and stress markers. Left: Scatter of loadings, colored by Training Module; Right upper: individual correlations of stress markers; Right lower: individual correlation of subfields; B). Multivariate associations between all subfields’ volume and intrinsic function and stress markers. Left: Scatter of loadings, colored by Training Module; Right upper: individual correlations of stress markers; Right lower: individual correlation of subfields.

Discussion

The goal of the current work was to investigate the effects of different types of mental training regimens on stress-related changes in the human hippocampus. The hippocampal formation is a highly plastic allocortex implicated in stress and emotional reactivity (14, 15, 19). In this study, we used automated segmentation to examine if the volumes of hippocampal subfields (SUB, CA1-3, CA4/DG) change in a large healthy sample over a 9-month longitudinal mental training study, the ReSource project (71). We investigated whether three different interoceptive and social mental Training Modules could lead to changes in hippocampal subfield volume. Additionally, we explored if these changes were associated with alterations in both intrinsic brain function and stress-related physiological changes, as indicated by shifts in diurnal cortisol levels due to training.

Earlier version: When comparing the differential efficiency of the 3-months mental Training Modules Presence, Affect, and Perspective against each other and a retest-control group on hippocampal subfield structure, we observed consistent increases in bilateral CA1-3 volume following socio-emotional Affect training relative to socio-cognitive Perspective training and no training in retest controls. Moreover, alterations in structure were mirrored by changes in functional connectivity of right CA1-3 following Affect versus Perspective training. In particular, we observed relative increases of functional connectivity between right CA1-3 and mPFC, and decreases between left CA1-3 and posterior insula, mainly driven by changes in connectivity following Perspective training. Evaluating training-related changes in diurnal cortisol output (cortisol awakening response, total diurnal output and diurnal slope), we observed that bilateral CA1-3 volume increases correlated with decreases in total diurnal cortisol output (assessed as the area under the curve with respect to ground, AUCg, sampled on 10 occasions over two consecutive days). Intrinsic connectivity of CA1-3 following Affect showed a positive association with left CA1-3 network change and diurnal slope and total diurnal cortisol output, where the latter associated with increased connectivity between left CA1-3 and parietal-occipital area. Interestingly, these associations were similar when combining Training Modules, suggesting the association between CA1-3 and diurnal cortisol markers is present irrespective of training content. Moreover, we additionally observed consistent associations between left CA1-3 and hair cortisol, a chronic stress marker, across trainings in a sub-sample of the current study. Last, performing multivariate analysis we again observed associations between CA1-3 volume and function plasticity and stress change, strongest in Affect. Yet combining all subfields in a single model indicated that other subfields also link to stress alterations, indicating that ultimately circuit-level alterations within the hippocampal formation relate to latent changes in diurnal stress markers across Training Modules.

Revised version: When comparing the differential efficiency of the 3-months mental Training Modules Presence, Affect, and Perspective against each other and a retest-control group on hippocampal subfield structure, we observed consistent increases in bilateral CA1-3 volume following socio-emotional Affect training relative to socio-cognitive Perspective training and no training in retest controls. Moreover, alterations in structure were mirrored by changes in functional connectivity of right CA1-3 following Affect versus Perspective training. In particular, we observed relative increases of functional connectivity between right CA1-3 and mPFC, and decreases between left CA1-3 and posterior insula, mainly driven by changes in connectivity following Perspective training. Evaluating training-related changes in diurnal cortisol output (cortisol awakening response, total diurnal output and diurnal slope), we observed that bilateral CA1-3 volume increases correlated with decreases in total diurnal cortisol output (assessed as the area under the curve with respect to ground, AUCg, sampled on 10 occasions over two consecutive days). Intrinsic connectivity of CA1-3 following Affect showed a positive association with left CA1-3 network change and diurnal slope and total diurnal cortisol output, where the latter associated with increased connectivity between left CA1-3 and parietal-occipital area. Interestingly, these associations were similar when combining Training Modules, suggesting the association between CA1-3 and diurnal cortisol markers is present irrespective of training content. Moreover, we additionally observed consistent associations between left CA1-3 and hair cortisol, a chronic stress marker, across trainings in a sub-sample of the current study. Finally, through conducting multivariate analysis, we once more noticed associations between changes in CA1-3 volume and functional adaptability and alterations in stress levels, particularly prominent within the Affect Module. Integrating all subfields into a unified model highlighted a distinct significance of CA1-3, although for the left hemisphere, we observed a more diverse range of contributions across subfields. In summary, we establish a connection between a socioemotional behavioral intervention, shifts in hippocampal subfield structure and function, and decreases in cortisol levels among healthy adults.

Our longitudinal, multi-modal approach could thus show that compassion-based mental training alters CA1-3 structure. Second, training-based increases in CA1-3 volume related to decreases in total diurnal cortisol release, suggesting that mental training and CA1-3 volume changes are linked to cortisol release. The results regarding changes in functional connectivity profiles were not as straightforward, but they did reveal a distinction between Affect and Perspective mental training in the CA1-3 region. Moreover, CA1-3 intrinsic functional change was associated with changes in diurnal cortisol slope and release, and long-term cortisol exposure. While the experimental nature of our training study allows concluding that CA1-3 structure changed as a function of Affect training, and that individual differences in CA1-3 structural change corresponded to cortisol release change, we cannot make any claims about which training-induced change caused the other. Thus, it is possible that, owing to the Affect module, the activation of emotion/motivation-related functional processes is key to reducing the daily stress load and associated cortisol release (79, 80). Such reduction in cortisol levels may then explain the observed downstream brain alterations. According to this interpretation, changes in CA1-3 volume may come secondary to stress reduction and consequently alterations in cortisol release following compassion training. Alternatively, emotion/compassion training may specifically targets CA volume and function, and, as per its role as the central break of the HPA axis, improves its capacity to inhibit cortisol release. This explanation could explain the lack of average diurnal cortisol (i.e., AUCg) change following Affect training per se (6), as it may be relevant for individual variations in brain change and thus be more difficult to detect based only on average change per module. In sum, it is likely that observed alterations in hippocampal structure and function, as well as their associations with diurnal cortisol change, are not explained by a single mechanism, but rather result as a combination of different factors. This interpretation is also supported by our multivariate observations. For example, given the anatomical and functional complexity of the hippocampal formation (26, 33) as well as the multifaceted cognitive processes underlying stress, it seems plausible that our observations are an emergent effect of multiple, cognitively distinct, functional sub-processes. Thus, future studies may directly test the potential specificity of the interrelationship between stress on the hippocampus using further targeted measurements.

The observed increases in CA1-3 volumes following socio-affective mental training were small-sized effects. However, findings were consistent when independently assessing the left and right hippocampus subfields. In particular, we observed that increases in CA1-3 volume after Affect training corresponded to a decrease in total diurnal cortisol as well as hair cortisol output. These results can be interpreted in line with the mainly inhibitory role of the hippocampus in stress regulation (20–22, 81). Specifically, the hippocampus is involved in the negative feedback inhibition of the HPA axis. Mineral- and glucocorticoid receptors are present in abundance in hippocampal neurons, from where they transmit the negative feedback signals of a wide range of glucocorticoid levels on HPA axis activity (22). The extremely high numbers of mineral- and glucocorticoid receptors make the hippocampus a prominent target for the neuro-toxic effects of glucocorticoids (82–84). In particular the CA1 may be susceptive to stress-based environmental effects due to synaptogenesis associated with NR2B subunits of glutamate receptors (NMDAR)(85). Along these lines, sustained exposure to high glucocorticoid levels was shown to relate to calcium influx, and may produce CA3 pyramidal neuronal damage, which has been reported in rodents and tree shrews (86–88). Next to demonstrating a consistent relationship between total daily cortisol output and hippocampal structure, the absence of univariate findings for cortisol awakening response (CAR), diurnal slope or hair cortisone levels may a divergence in the sensitivity of alternative cortisol-based stress markers to structural neuroimaging markers. It is of note that the current work relies on a segmentation approach of hippocampal subfields including projection to MNI template space, an implicit correction for total brain volume through the use of a stereotaxic reference frame. Some caution for this method may be warranted, as complex hippocampal anatomy can in some cases lead to over-as well as underestimation of subfield volumes, as well as subfield boundaries may not always be clearly demarcated (89). Future work, studying the hippocampal surface at higher granularity, for example though unfolding the hippocampal sheet (35, 76, 90, 91), may further help with both alignment and identification of not only subfield-specific change but also alterations as a function of the hippocampal long axis, a key dimension of hippocampal structural and functional variation that was not assessed in the current work (33, 92).

Structural MRI findings were complemented by the separate assessment of task-free (“resting-state”) functional connectivity networks. Identification of networks that show coupled spontaneous brain activity through resting-state fMRI is currently considered an effective approach to study whole-brain functional connectivity (93–95). In the current cohort, we could demonstrate widespread patterns of hippocampal functional connectivity to mesiotemporal, lateral temporal, together with anterior as well as posterior midline regions, lateral temporo-parietal, and dorsolateral prefrontal cortices - a pattern in excellent accordance to previous studies probing hippocampal functional connectivity at rest in healthy populations (28, 35, 38, 96, 97), and outlining “mesiotemporal” components of default-mode networks (98, 99). Assessing modulations of connectivity by mental training, we could provide independent, yet weak, support for a specific relationship of the socio-affective Affect training, relative to socio-cognitive Perspective Training Modules, on hippocampal network embedding. In particular, we observed an increased functional integration of the right CA1-3 with medial prefrontal cortical regions (mPFC) in individuals following Affect relative to Perspective training. Studies in rats and non-human primates have demonstrated a high density of glucocorticoid receptors in the mPFC (100, 101). Accordingly, the mPFC, like the hippocampus, was shown to play a key role in HPA-axis regulation (50, 81, 102, 103). In a previous positron emission tomography study, glucose metabolism in the mPFC was negatively associated with acute stress-induced salivary cortisol increases; notably, the authors observed a negative metabolic coupling between mPFC areas and the mesiotemporal lobe (104). In related work on isocortical changes in structure and intrinsic function following the ReSource training, we have observed structural changes in insular, opercular and orbitofrontal regions following Affect training (9, 73). At the same time, we observed little change in large-scale functional organization, relative to changes observed following Presence and Perspective training. Previous work has implicated the hippocampal formation at the nexus of multiple large-scale networks and cortical organization (28, 92). Indeed, it may be that particular changes in the CA1-3 are central in coordinating the signal flow within the hippocampal complex, coordinating the balance between large-scale association networks in the isocortex (28). Integrating this with our empirical observation of Affect training taking up a regulatory or stabilizing functional role, relative to Perspective and Presence training, it is possible that such alterations are orchestrated by adaptive processes (105). Future work may be able to further disentangle the causal relationship between iso- and allo-cortical structure and function, and the role of specific hippocampal subfields.

Using univariate approaches, we could observe that training-induced HC volume increases following socio-affective mental training overlapped with reductions in cumulative diurnal cortisol release. Additionally, we observed functional connectivity decrease between left CA1-3 and parietal-occipital area in individuals showing reduced diurnal cortisol release and overall connectivity decreases of left CA1-3 relating to reductions in diurnal cortisol slope. Importantly, these associations could be found also when including Presence and Perspective in our analysis, suggesting of a domain-general relationship between diurnal cortisol alterations and CA1-3 volume and function. In line with our observations in univariate analysis, we found multivariate associations between hippocampal subfield volume, intrinsic function and cortisol markers. Again, the contribution of volume and intrinsic function was inverse. This may possibly relate to the averaging procedure of the functional networks. Combined, outcomes of our univariate and multivariate analyses point to an association between change in hippocampal subfields and stress markers, and that these changes, at the level of the individual, ultimately reflect complex interactions within and across hippocampal subfields and may capture different aspects of diurnal stress. Future work may more comprehensively study the plasticity of the hippocampal structure, and link this to intrinsic functional change and cortisol to gain full insights in the specificity and system-level interplay across subfields, for example using more detailed hippocampal models (90). Incorporating further multivariate, computational, models is needed to further unpack and investigate the complex and nuanced association between hippocampal structure and function, in particular in relation to subfield plasticity and short and long-term stress markers. In line with our multivariate observations, in other work from the ReSource study we observed mixed specificity of stress-reducing effects as a function of mental training. For example, both social modules, that is the Affect and Perspective trainings, reduced acute cortisol reactivity to a psychosocial stressor (106), which is considered a dynamic state of HPA axis activity(7). Regarding the CAR, only Affect training was able to reduce this dynamic cortisol response to awakening, known to reflect anticipatory stress (6). Lastly, regarding hair cortisol, a long-term measure of systemic stress, all Training Modules were shown to be equally effective in stress reduction over a training period of three to nine months (5). In our work we observed a consistent association between left CA1-3 volume and functional increases and hair cortisol decreases, hinting at a potential relationship between CA1-3 and both short-term and long-term stress level changes.

Earlier version: Overall, different types of mental training result in stress reduction (e.g. (5–7)). In a recent paper we argue that the variable pattern of mental training effects on different cortisol indices may be explained by the functional roles of these indices (6). Thus, indices reflecting dynamic HPA axis properties, such as acute stress reactivity and the CAR, were suggested to change with Affect and Perspective training (also see (7)). Hair cortisol as a marker of cumulative stress load likely reflecting the low-grade and continuous strain inherent to daily hassles (107–109), was contrarily suggested to change independent of training type (also see (5)). The current findings do not necessarily contradict this reasoning, due to differences in interpretation of group-level and individual-level changes. Although we observed that CA1-3 volume was selectively increased by socio-affective mental training at the group level, and that individual differences in CA1-3 volume increase within the Affect module correlated with reduced diurnal cortisol release, the pattern linking bilateral CA1-3 volume increases with reduced diurnal cortisol release was also present when all modules were combined. Similarly, in follow-up analysis on functional alterations of hippocampal subfields, we could observe group-level increases in connectivity to mPFC for right, but not left, CA1-3, when comparing socio-affective and socio-cogni tive training. Though right CA1-3 group-level changes did not link to individual level change in cortisol markers following Affect training, individual level changes in left CA1-3 corresponded to changes in cortisol markers, again following Affect but also across all practices combined. Thus, we cannot at this point derive a consistent pattern of how mental training influences different indices of cortisol activity, yet we do find a consistent change in CA1-3 following Affect training, and observe domain-general patterns of change associations between CA1-3 and cortisol markers, indicating CA1-3 may play a central role within the context of Affect training and diurnal stress reduction based on univariate analysis. Our multivariate analysis further nuanced this observation, but again pointed to an overall association between hippocampal subfield changes and cortisol changes, but this time more at a systems level. Lastly, from a mechanistic viewpoint, we hypothesize that Affect training stimulates emotion-motivational (reward) systems associated with positive affect (79, 80), and regulated by oxytocin and opiates (110, 111). Since these neuropeptides are also involved in stress regulation (112, 113), they could be considered to provide a double hit, and prime candidates to mediate hippocampal volume increase and stress reduction in particular following compassion-based practice, yet also present following other practices.

Revised version: Overall, different types of mental training result in stress reduction (e.g. (5–7)). In a recent paper we argue that the variable pattern of mental training effects on different cortisol indices may be explained by the functional roles of these indices (6). Thus, indices reflecting dynamic HPA axis properties, such as acute stress reactivity and the CAR, were suggested to change with Affect and Perspective training (also see (7)). Hair cortisol as a marker of cumulative stress load likely reflecting the low-grade and continuous strain inherent to daily hassles (107–109), was contrarily suggested to change independent of training type (also see (5)). The current findings do not necessarily contradict this reasoning, due to differences in interpretation of group-level and individual-level changes. Although we observed that CA1-3 volume was selectively increased by socio-affective mental training at the group level, and that individual differences in CA1-3 volume increase within the Affect module correlated with reduced diurnal cortisol release, the pattern linking bilateral CA1-3 volume increases with reduced diurnal cortisol release was also present when all modules were combined. Similarly, in follow-up analysis on functional alterations of hippocampal subfields, we could observe group-level increases in connectivity to mPFC for right, but not left, CA1-3, when comparing socio-affective and socio-cogni tive training. Though right CA1-3 group-level changes did not link to individual level change in cortisol markers following Affect training, individual level changes in left CA1-3 corresponded to changes in cortisol markers, again following Affect but also across all practices combined. Thus, we cannot at this point derive a consistent pattern of how mental training influences different indices of cortisol activity, yet we do find a consistent change in CA1-3 following Affect training, and observe domain-general patterns of change associations between CA1-3 and cortisol markers, indicating CA1-3 may play a central role within the context of Affect training and diurnal stress reduction based on univariate analysis. Although the univariate examination of changes specific to modules in volume and connections within the Affect Module presents how changes in cortisol align with group-level rises in CA1-3 volume, the multivariate analysis extended this observation through considering individual-level alterations not discernible at the group level through a data-driven method. These results generally corresponded with observations at the group level but offer additional insights into specificity, and hint at system-level alterations. Lastly, from a mechanistic viewpoint, we hypothesize that Affect training stimulates emotion-motivational (reward) systems associated with positive affect (79, 80), and regulated by oxytocin and opiates (110, 111). Since these neuropeptides are also involved in stress regulation (112, 113), they could be considered to provide a double hit, and prime candidates to mediate hippocampal volume increase and stress reduction in particular following compassion-based practice, yet also present following other practices.

It is of note that non-adherence to saliva sampling in ambulatory settings has been shown to exert a significant impact on the resulting cortisol data (114, 115) and that the present data does not fully conform to the recently provided consensus guidelines on the assessment of the CAR (116, 117), which were published after the conception of our study. Most importantly, we did not employ objective measures for the verification of participants’ sampling times. Hence, diurnal cortisol data have to be treated with some caution since the possibility of non-adherence-related confounding cannot be excluded (114–117). We nevertheless addressed the issue of non-adherence through an experience sampling approach based on mobile phones handed out to our participants. As shown by the relatively low proportion of missing data, these devices may have boosted adherence by reminding participants of a forthcoming sampling time-point.

To conclude, using a longitudinal model, we investigated how different types of mental training differentially result in changes in hippocampal subfield volume, resting-state functional networks, and stress-related markers of diurnal cortisol and hair cortisol. We find that only the 3-months Affect training module cultivating compassion and care, rather than attention-based (Presence) or socio-cognitive (Perspective) training, related to an increase of hippocampal CA1-3 subfield volume, with corresponding alterations in functional connectivity and a reduction in total diurnal cortisol output. Across analyses we observed consistent alterations between cortisol change and CA1-3 volume and function, pinpointing this region as a potential target for further investigations on stress and the human brain. Lastly, our multivariate analyses also point to a circuit level understanding of latent diurnal stress scores. Our results may be informative for the development of targeted interventions to reduce stress, and inspire the update of models on the role of different hippocampal formations for human socio-emotional and stress-related processes.

Methods

The specifics on the experimental design are the similar to related works in the same sample (9, 118). They are provided again here for completeness.

Participants

Earlier version: We recruited a total of 332 healthy adults (197 women, mean±SD=40.7±9.2 years, 20-55 years), in the winters of 2012/2013 and 2013/2014. Participant eligibility was determined through a multi-stage procedure that involved several screening and mental health questionnaires, together with a phone interview [for details, see (71)]. Subsequently, a face-to-face mental health diagnostic interview with a trained clinical psychologist was carried out. The interview included a computer-assisted German version of the Structured Clinical Interview for DSM-IV Axis-I disorders, SCID-I DIA-X (118), and a personal interview, SCID-II, for Axis-II disorders (119, 120). Participants were excluded if they fulfilled criteria for: i) an Axis-I disorder within the past two years, ii) schizophrenia, psychotic disorders, bipolar disorder, or substance dependency, or iii) an Axis-II disorder at any time in their life. Participants taking medication influencing the HPA axis were also excluded. None of the participants had a history of suffering from neurological disorders or head trauma, based on an in-house self-report questionnaire completed prior to the neuroimaging investigations. Included participants furthermore underwent a diagnostic radiological evaluation to rule out the presence of mass lesions (e.g., tumors, vascular malformations). The study was approved by the Research Ethics Committees of University of Leipzig (#376/12-ff) and Humboldt University in Berlin (#2013-02, 2013-29, 2014-10), and all participants provided written informed consent prior to participation. The study was registered with the Protocol Registration System of ClinicalTrials.gov under the title “Plasticity of the Compassionate Brain” with the Identifier: NCT01833104. For more details on recruiting and sample selection, please see (71).

Revised version: We recruited a total of 332 healthy adults (197 women, mean±SD=40.7±9.2 years, 20-55 years), in the winters of 2012/2013 and 2013/2014. Participant eligibility was determined through a multi-stage procedure that involved several screening and mental health questionnaires, together with a phone interview [for details, see (71)]. Subsequently, a face-to-face mental health diagnostic interview with a trained clinical psychologist was carried out. The interview included a computer-assisted German version of the Structured Clinical Interview for DSM-IV Axis-I disorders, SCID-I DIA-X (119), and a personal interview, SCID-II, for Axis-II disorders (120, 121). Participants were excluded if they fulfilled criteria for: i) an Axis-I disorder within the past two years, ii) schizophrenia, psychotic disorders, bipolar disorder, or substance dependency, or iii) an Axis-II disorder at any time in their life. Participants taking medication influencing the HPA axis were also excluded. None of the participants had a history of suffering from neurological disorders or head trauma, based on an in-house self-report questionnaire completed prior to the neuroimaging investigations. Included participants furthermore underwent a diagnostic radiological evaluation to rule out the presence of mass lesions (e.g., tumors, vascular malformations). The study was approved by the Research Ethics Committees of University of Leipzig (#376/12-ff) and Humboldt University in Berlin (#2013-02, 2013-29, 2014-10), and all participants provided written informed consent prior to participation. The study was registered with the Protocol Registration System of ClinicalTrials.gov under the title “Plasticity of the Compassionate Brain” with the Identifier: NCT01833104. For more details on recruiting and sample selection, please see (71).

Sample size estimation and group allocation

Overall, 2595 people signed up for the ReSource study in winter 2012/2013. Of these individuals, 311 potential participants met all eligibility criteria. From the latter group, 198 were randomly selected as the final sample. Participants were selected from the larger pool of potential participants and assigned to cohorts using bootstrapping without replacement, creating cohorts that did not differ (omnibus test p<0.1) in demographics (age, gender, marital status, income, and IQ) or self-reported traits (depression, empathy, interoceptive awareness, stress level, compassion for self and others, alexithymia, general mental health, anxiety, agreeableness, conscientiousness, extraversion, neuroticism, and openness). Seven participants dropped out of the study after assignment but before data collection began, leaving 30 participants in RCC1, 80 in TC1, and 81 in TC2.

2144 people applied for the second wave of the study in winter 2013/2014. Of these people, 248 potential participants met all the eligibility criteria. From the latter pool, 164 were then randomly selected as the final sample. Participants were selected from the larger pool of potential participants and assigned to cohorts using bootstrapping without replacement, creating cohorts that did not differ significantly (omnibus test, p>0.1) from the Winter 2012/2013 cohorts or from one another in demographics (age, gender, marital status, income, and IQ) or self-reported traits (depression, empathy, interoceptive awareness, stress level, compassion for self and others, alexithymia, general mental health, anxiety, agreeableness, conscientiousness, neuroticism, and openness). The control cohorts (RCC1, RCC2, and RCC1&2) were significantly lower in extraversion than TC3; participants in the control cohorts were also more likely to have children than participants in TC3. Twenty-three participants dropped out of the study after assignment but before data collection began, leaving 81 participants in TC3 and 60 in RCC2. See further (71).

ReSource training program

In the ReSource Project, we investigated the specific effects of commonly used mental training techniques by parceling the training program into three separate modules (Presence, Affect and Perspective). Participants were selected from a larger pool of potential volunteers by bootstrapping without replacement, creating cohorts not differing significantly with respect to several demographic and self-report traits (71). Each cultivated distinct cognitive and socio-affective capacities (69). Participants were divided in two 9-month training cohorts experiencing the modules in different orders, one 3-month Affect training cohort and one retest control cohort. In detail, two training cohorts (TC1, TC2) started their training with the mindfulness-based Presence module. They then underwent Affect and Perspective modules in different orders thereby acting as mutual active control groups. To isolate the specific effects of the Presence module, a third training cohort (TC3) underwent the 3-month Affect module only (Fig. 1B).

As illustrated in Fig 1A, the core psychological processes targeted in the Presence module are attention and interoceptive awareness, which are trained through the two meditation-based core exercises Breathing Meditation and Body Scan. The Affect module targets the cultivation of social emotions such as compassion, loving kindness and gratitude. It also aims to enhance prosocial motivation and dealing with difficult emotions. The two core exercises of the Affect module are Loving-kindness Meditation and Affect Dyad. In the Perspective module participants train meta-cognition and perspective-taking on self and others through the two core exercises Observing-thoughts Meditation and Perspective Dyad. The distinction between Affect and Perspective modules reflects research identifying distinct neural routes to social understanding: One socio-affective route including emotions such as empathy and compassion, and one socio-cognitive route including the capacity to mentalize and take perspective on self and others (for details on the scientific backbone of this division see: (71)).

Earlier version: The two contemplative dyads are partner exercises that were developed for the ReSource training (121). They address different skills such as perspective taking on self and others (Perspective dyad) or gratitude, acceptance of difficult emotions and empathic listening (Affect dyad), but are similar in structure (for details see: (71)). In each 10-min dyadic practice, two randomly paired participants share their experiences with alternating roles of speaker and listener. The dyadic format is designed to foster interconnectedness by providing opportunities for self-disclosure and non-judgmental listening (71, 121). Our recommendation was to train for a minimum of 30 minutes (e.g. 10 minutes contemplative dyad, 20 minutes classic meditation) on five days per week.

Revised version: The two contemplative dyads are partner exercises that were developed for the ReSource training (122). They address different skills such as perspective taking on self and others (Perspective dyad) or gratitude, acceptance of difficult emotions and empathic listening (Affect dyad), but are similar in structure (for details see: (71)). In each 10-min dyadic practice, two randomly paired participants share their experiences with alternating roles of speaker and listener. The dyadic format is designed to foster interconnectedness by providing opportunities for self-disclosure and non-judgmental listening (71, 122). Our recommendation was to train for a minimum of 30 minutes (e.g. 10 minutes contemplative dyad, 20 minutes classic meditation) on five days per week.

MRI acquisition

MRI data were acquired on a 3T Siemens Magnetom Verio (Siemens Healthcare, Erlangen, Germany) using a 32-channel head coil. Structural images were acquired using a T1-weighted 3D-MPRAGE sequence (repetition time [TR]=2300 ms, echo time [TE]=2.98 ms, inversion time [TI]=900 ms, flip angle=7°; 176 sagittal slices with 1mm slice thickness, field of view [FOV]=240×256 mm2, matrix=240×256, 1×1×1 mm3 voxels). We recorded task-free functional MRI using a T2*-weighted gradient EPI sequence (TR=2000ms, TE=27ms, flip angle=90°; 37 slices tilted at approximately 30° with 3 mm slice thickness, FOV=210×210mm2, matrix=70×70, 3×3×3 mm3 voxels, 1 mm gap; 210 volumes per session). During the functional session, participants were instructed to lie still in the scanner, think of nothing in particular, and fixate a white cross in the center of a black screen.

Structural MRI analysis: Hippocampal subfield volumetry

Earlier version: Based on the available high-resolution T1-weighted images subiculum (SUB), CA1-3, and CA4/DG were segmented using a patch-based algorithm in all participants individually (see further (74)). Shortly, this procedure uses a population-based patch normalization relative to a template library (122), providing good time and space complexity. In previous validations work, this algorithm has shown high segmentation accuracy of hippocampal subfields (74), and in detecting hippocampal subfield pathology in patients with epilepsy (91). It was furthermore demonstrated that these representations can be used to probe sub-regional functional organization of the hippocampus (35, 36). Hippocampal volumes were estimated based on T1w data that were linearly registered to MNI152 using FSL flirt (http://www.fmrib.ox.ac.uk/fsl/), such that intracranial volume was implicitly controlled for.

Revised version: Based on the available high-resolution T1-weighted images subiculum (SUB), CA1-3, and CA4/DG were segmented using a patch-based algorithm in all participants individually (see further (74)). Shortly, this procedure uses a population-based patch normalization relative to a template library (123), providing good time and space complexity. In previous validations work, this algorithm has shown high segmentation accuracy of hippocampal subfields (74), and in detecting hippocampal subfield pathology in patients with epilepsy (91). It was furthermore demonstrated that these representations can be used to probe sub-regional functional organization of the hippocampus (35, 36). Hippocampal volumes were estimated based on T1w data that were linearly registered to MNI152 using FSL flirt (http://www.fmrib.ox.ac.uk/fsl/), such that intracranial volume was implicitly controlled for.

Earlier version: As previously reported (123), for successful hippocampus segmentations, an initial quality check was conducted by two independent raters, R.L. and L.P. Both raters were blind to participant characteristics including age, sex, and training or control group. In short, each segmentation was rated for quality on a scale of 1–10, with points being subtracted depending on the severity of detected flaws. One point was subtracted for minor flaws, e.g. part of a segmentation extends slightly beyond the hippocampal boundary, or does not cover a small aspect of the hippocampal formation. Two points were subtracted for medium flaws, e.g. gaps between sub-field segmentations. Finally, major flaws immediately qualified for resampling, and included e.g. one or more subfield segmentations being clearly misplaced. Given a minimum of 70% inter-rater reliability, segmentation ratings were then averaged and evaluated, with scores of 5 and lower qualifying for reprocessing with the algorithm. Following this second round of processing, segmentations were rated again. Any remaining segmentations with average scores lower than 5 were excluded from the analysis.

Revised version: As previously reported (124), for successful hippocampus segmentations, an initial quality check was conducted by two independent raters, R.L. and L.P. Both raters were blind to participant characteristics including age, sex, and training or control group. In short, each segmentation was rated for quality on a scale of 1–10, with points being subtracted depending on the severity of detected flaws. One point was subtracted for minor flaws, e.g. part of a segmentation extends slightly beyond the hippocampal boundary, or does not cover a small aspect of the hippocampal formation. Two points were subtracted for medium flaws, e.g. gaps between sub-field segmentations. Finally, major flaws immediately qualified for resampling, and included e.g. one or more subfield segmentations being clearly misplaced. Given a minimum of 70% inter-rater reliability, segmentation ratings were then averaged and evaluated, with scores of 5 and lower qualifying for reprocessing with the algorithm. Following this second round of processing, segmentations were rated again. Any remaining segmentations with average scores lower than 5 were excluded from the analysis.

Task-free functional MRI analysis: Hippocampal connectivity

Earlier version: Processing was based on DPARSF/REST for Matlab [http://www.restfmri.net (124)]. We discarded the first 5 volumes to ensure steady-state magnetization, performed slice-time correction, motion correction and realignment, and co-registered functional time series of a given subject to the corresponding T1-weighted MRI. Images underwent unified segmentation and registration to MNI152, followed by nuisance covariate regression to remove effects of average WM and CSF signal, as well as 6 motion parameters (3 translations, 3 rotations). We included a scrubbing (125) that modeled time points with a frame-wise displacement of ≥0.5 mm, together with the preceding and subsequent time points as separate regressors during nuisance covariate correction.

Revised version: Processing was based on DPARSF/REST for Matlab [http://www.restfmri.net (125)]. We discarded the first 5 volumes to ensure steady-state magnetization, performed slice-time correction, motion correction and realignment, and co-registered functional time series of a given subject to the corresponding T1-weighted MRI. Images underwent unified segmentation and registration to MNI152, followed by nuisance covariate regression to remove effects of average WM and CSF signal, as well as 6 motion parameters (3 translations, 3 rotations). We included a scrubbing (126) that modeled time points with a frame-wise displacement of ≥0.5 mm, together with the preceding and subsequent time points as separate regressors during nuisance covariate correction.

We linearly co-registered the extracted hippocampal subfield volumes with the functional MRI data for each individual using FSL flirt (http://www.fmrib.ox.ac.uk/fsl/), followed by nearest neighbor interpolation. Following, we generated functional connectivity maps from both the left and right hippocampal subfields in each individual. Functional connectivity was calculated as the correlation between the mean time series of the seed region and the time series of all cortical parcels based on the Schaefer 400 parcellation. To render them normally distributed and scale the profiles across participants, correlation coefficients underwent a Fisher r-to-z transformation and were rescaled, resulting in connectivity profiles between 0 and 1 for each participant and timepoint. Functional networks were defined as the top 10% regions based on mean connectivity profile of the respective subfield in the ipsilateral hemisphere at baseline. Individuals with a framewise-displacement of >0.3mm (<5%) were excluded.

Diurnal cortisol assessments

Earlier version: For cortisol assessment, 14 saliva samples (7 per day) were obtained over the course of two consecutive weekdays (Mondays/Tuesdays, Wednesdays/Thursdays or Thursdays/Fridays, depending on participant availability). In detail, samples were taken upon free awakening (while still in bed; S1) and at 30 minutes, 60 minutes, 4, 6, 8 and 10 hours after awakening. Saliva was collected using Salivette collection devices (Sarstedt, Nuembrecht, Germany). Participants were instructed to place collection swabs in their mouths and to refrain from chewing for 2 minutes. They were asked to not eat, drink (except water), or brush their teeth during the 10 minutes before sampling, and to not smoke during the 30 minutes before sampling. If deviating from this guideline, they were asked to thoroughly rinse their mouth with water before taking a sample. Participants otherwise followed their normal daily routine. To maximize adherence to the sampling protocol, participants were given pre-programmed mobile devices using an inhouse application that reminded them to take each (except the first) Salivette at the designated time. Sampling times of the non-morning probes were jittered (+/-15 min) to avoid complete predictability. Samples were kept in the freezer until returned to the laboratory, where they were stored at -30 °C until assay (at the Department of Biological and Clinical Psychology, University of Trier, Germany). Cortisol levels (expressed in nmol/l) were determined using a time-resolved fluorescence immunoassay (126) with intra-/inter-assay variability of 10/12%.

Revised version: For cortisol assessment, 14 saliva samples (7 per day) were obtained over the course of two consecutive weekdays (Mondays/Tuesdays, Wednesdays/Thursdays or Thursdays/Fridays, depending on participant availability). In detail, samples were taken upon free awakening (while still in bed; S1) and at 30 minutes, 60 minutes, 4, 6, 8 and 10 hours after awakening. Saliva was collected using Salivette collection devices (Sarstedt, Nuembrecht, Germany). Participants were instructed to place collection swabs in their mouths and to refrain from chewing for 2 minutes. They were asked to not eat, drink (except water), or brush their teeth during the 10 minutes before sampling, and to not smoke during the 30 minutes before sampling. If deviating from this guideline, they were asked to thoroughly rinse their mouth with water before taking a sample. Participants otherwise followed their normal daily routine. To maximize adherence to the sampling protocol, participants were given pre-programmed mobile devices using an inhouse application that reminded them to take each (except the first) Salivette at the designated time. Sampling times of the non-morning probes were jittered (+/-15 min) to avoid complete predictability. Samples were kept in the freezer until returned to the laboratory, where they were stored at -30 °C until assay (at the Department of Biological and Clinical Psychology, University of Trier, Germany). Cortisol levels (expressed in nmol/l) were determined using a time-resolved fluorescence immunoassay (127) with intra-/inter-assay variability of 10/12%.

Earlier version: Raw cortisol data were each treated with a natural log transformation to remedy skewed distributions. Across the full sample, any values diverging more than 3 SD from the mean were labeled outliers and winsorized to the respective upper or lower 3 SD boundary to avoid influential cases. Logged and winsorized cortisol data was then averaged across the two sampling days, and the most commonly used summary indices of diurnal cortisol activity were calculated (13). The CAR was quantified as a change score from S1 to either the 30-or 60-minute post-awakening sample, depending on the individual peak in hormone levels. If participants peaked at S1 rather than at 30 or 60 minutes thereafter, the 30-minute data point was used to operationalize the (inverse) CAR, given that it was always closer in magnitude to S1 than the 60-minute data point. The cortisol decline over the course of the day (diurnal slope) was operationalized as a change score from baseline to the final sample of the day (at 600 minutes after awakening). Total daily cortisol output was operationalized as the area under the curve with respect to ground, AUCg (127), which considers the difference between the measurements from each other (i.e., the change over time) and the distance of these measures from zero (i.e., the level at which the change over time occurs). Awakening, 240, 360, 480, and 600 minutes post-awakening cortisol values were included in the calculation of the AUCg. To prevent it from having an undue influence, the CAR samples at 30 and 60 minutes were excluded from the total output score calculation. On each sampling day, awakening time and sleep duration were registered using the pre-programmed mobile device immediately upon awakening in parallel to taking the first Salivette. These measures were averaged across the two sampling days to minimize situational influences.

Revised version: Raw cortisol data were each treated with a natural log transformation to remedy skewed distributions. Across the full sample, any values diverging more than 3 SD from the mean were labeled outliers and winsorized to the respective upper or lower 3 SD boundary to avoid influential cases. Logged and winsorized cortisol data was then averaged across the two sampling days, and the most commonly used summary indices of diurnal cortisol activity were calculated (13). The CAR was quantified as a change score from S1 to either the 30-or 60-minute post-awakening sample, depending on the individual peak in hormone levels. If participants peaked at S1 rather than at 30 or 60 minutes thereafter, the 30-minute data point was used to operationalize the (inverse) CAR, given that it was always closer in magnitude to S1 than the 60-minute data point. The cortisol decline over the course of the day (diurnal slope) was operationalized as a change score from baseline to the final sample of the day (at 600 minutes after awakening). Total daily cortisol output was operationalized as the area under the curve with respect to ground, AUCg (128), which considers the difference between the measurements from each other (i.e., the change over time) and the distance of these measures from zero (i.e., the level at which the change over time occurs). Awakening, 240, 360, 480, and 600 minutes post-awakening cortisol values were included in the calculation of the AUCg. To prevent it from having an undue influence, the CAR samples at 30 and 60 minutes were excluded from the total output score calculation. On each sampling day, awakening time and sleep duration were registered using the pre-programmed mobile device immediately upon awakening in parallel to taking the first Salivette. These measures were averaged across the two sampling days to minimize situational influences.

Assay of Steroid Hormone Concentration in Hair

Earlier version: Please see further details on sample and dropout in (5). To evaluate cortisol and cortisone, hair strands were taken as close as possible to the scalp from a posterior vertex position at T0 and after each following timepoint (T0-T3). Hair samples were enfolded in aluminum foil and stored in the dark at room temperature until assay at the Department of Psychology, TU Dresden, Germany. We evaluated the proximal 3-cm segment of hair to study accumulation of cortisol and cortisone over each 3-month period, based on the assumption of an average hair growth rate of 1 cm/month (128). Hormone concentrations were captured using liquid chromatography–tandem mass spectrometry, the current criterion standard approach for hair steroid analysis (129). All hormone concentrations were reported in picograms per milligram. For the current longitudinal research aim, all samples of one participant were always run with the same reagent batch to avoid intraindividual variance due to batch effects.

Revised version: Please see further details on sample and dropout in (5). To evaluate cortisol and cortisone, hair strands were taken as close as possible to the scalp from a posterior vertex position at T0 and after each following timepoint (T0-T3). Hair samples were enfolded in aluminum foil and stored in the dark at room temperature until assay at the Department of Psychology, TU Dresden, Germany. We evaluated the proximal 3-cm segment of hair to study accumulation of cortisol and cortisone over each 3-month period, based on the assumption of an average hair growth rate of 1 cm/month (129). Hormone concentrations were captured using liquid chromatography–tandem mass spectrometry, the current criterion standard approach for hair steroid analysis (130). All hormone concentrations were reported in picograms per milligram. For the current longitudinal research aim, all samples of one participant were always run with the same reagent batch to avoid intraindividual variance due to batch effects.

Quality control and case selection

Structural MRI data without artifacts and acceptable automated segmentations were available in 943 participants. Functional MRI data were available in 849 participants. We opted to have consistent sample sizes in structure and function and therefor including only people that had both structural and functional data available. Please see Table 4. for participant numbers across timepoints and measures for structural and functional data.

Among those, salivary cortisol measures were available in Presence n= 85 (53 females, age= 40.87 std 9.69, 20-55), Affect n= 89 (50 females, age= 40.11 std 9.87, 20-55), Perspective n=81 (48 females, age= 40.14 std 9.78, 20-55). Hair cortisol change scores were available in Presence n= 31 (21 females, age= 39.55 std 10.40, 20-54), Affect n= 44 (24 females, age= 37.52 std 10.78, 20-54), Perspective n= 41 (24 females, age= 38.14 std 10.51, 20-54).

Statistical analyses

Earlier version: Using SurfStat (130, 131), we carried out structural and functional MRI analysis for the left and right hippocampal subfield difference scores between different 3-month timepoints. All models statistically corrected for nuisance effects of age and sex, as well as random effect of subject. Main contrasts considered in the group analyses concern Presence versus Active Control (T0-T1) and Affect versus Perspective (T1-T3). Additionally, investigations include analyses versus Retest Control Cohort as well as subgroups defined by training cohort and timepoint. In case of multiple comparison, we performed Bonferroni correction (132).

Revised version: Using SurfStat (131, 132), we carried out structural and functional MRI analysis for the left and right hippocampal subfield difference scores between different 3-month timepoints. All models statistically corrected for nuisance effects of age and sex, as well as random effect of subject. Main contrasts considered in the group analyses concern Presence versus Active Control (T0-T1) and Affect versus Perspective (T1-T3). Additionally, investigations include analyses versus Retest Control Cohort as well as subgroups defined by training cohort and timepoint. In case of multiple comparison, we performed Bonferroni correction (133).

Partial least squares analysis

To assess potential relationships between cortisol change and hippocampal subfield volume and functional change, we performed a partial least squares analysis (PLS) (77, 78). PLS is a multivariate associative model that to optimizes the covariance between two matrices, by generating latent components (LCs), which are optimal linear combinations of the original matrices (77, 78). In our study, we utilized PLS to analyze the relationships between change in volume and intrinsic function of hippocampal subfields and diurnal cortisol measures. Here we included all Training Modules and regressed out effects of age, sex, and random effects of subject on the brain measures before conducting the PLS analysis. The PLS process involves data normalization within training groups, cross-covariance, and singular value decomposition.

Subsequently, subfield and behavioral scores are computed, and permutation testing (1000 iterations) is conducted to evaluate the significance of each latent factor solution (FDR corrected). We report then the correlation of the individual hippocampal and cortisol markers with the latent factors. To estimate confidence intervals for these correlations, we applied a bootstrapping procedure that generated 100 samples with replacement from subjects’ RSFC and behavioral data.